---
title: "Q3"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(comment = "",echo = FALSE)
```
<style>
  body {
    text-align: justify
    }
</style>

### 1.

<b>Take a look at the 'iris' dataset that comes with R. The data can be loaded with the code:</b>

```{r w1,eval = FALSE,echo = TRUE}
library(datasets)
data(iris)
?iris
```

<b>There will be an object called 'iris' in your workspace. In this dataset, what
is the mean of 'Sepal.Length' for the species virginica? Please round your
answer to the nearest whole number.

(Only enter the numeric result and nothing else.)</b>

```{r q1a,exercise = TRUE}

```

```{r q1a-solution}
round(mean(iris$Sepal.Length[iris$Species == 'virginica'],na.rm = TRUE))
```

```{r q1a-check}
grade_this(if(identical(.result,7)) {
  pass("Great job!")
} else {
  fail("Start over again!")
})
```

### 2.

```{r q2}
question("Continuing with the 'iris' dataset from the previous question, what R
         code returns a vector of the means of the variable
         'Sepal.Length','Sepal.Width','Petal.Length' and 'Petal.Width'?",
  answer('apply(iris[, 1:4], 2, mean)',correct = TRUE),
  answer('colMeans(iris)'),
  answer('apply(iris[, 1:4], 1, mean)'),
  answer('rowMeans(iris[, 1:4])'),
  answer('apply(iris, 1, mean)'),
  answer('apply(iris, 2, mean)'))
```

### 3.
<b>Load the 'mtcars' dataset in R with the following code</b>

```{r w,eval = FALSE,echo = TRUE}
library(datasets)
data(mtcars)
?mtcars
```

```{r q3}
question('How can one calculate the average miles per gallon (mpg) by number of
         cylinders in the car (cyl)? Select all that apply.',
  answer('sapply(mtcars, cyl, mean)'),
  answer('mean(mtcars$mpg, mtcars$cyl)'),
  answer('sapply(split(mtcars$mpg, mtcars$cyl), mean)',correct = TRUE),
  answer('tapply(mtcars$mpg, mtcars$cyl, mean)',correct = TRUE),
  answer('with(mtcars, tapply(mpg, cyl, mean))',correct = TRUE),
  answer('lapply(mtcars, mean)'),
  answer('apply(mtcars, 2, mean)'),
  answer('split(mtcars, mtcars$cyl)'),
  answer('tapply(mtcars$cyl, mtcars$mpg, mean)'))
```


### 4.

<b>Continuing with the 'mtcars' dataset from the previous Question, what is the
absolute difference between the average horsepower of 4-cylinder cars and the
average horsepower of 8-cylinder cars?

(Please round your final answer to the nearest whole number. Only enter the
numeric result and nothing else.)</b>

```{r q4,exercise = TRUE}

```

```{r q4-solution}
round(abs(mean(mtcars$hp[mtcars$cyl == 4]) - mean(mtcars$hp[mtcars$cyl == 8])))
```

```{r q4-check}
grade_this(if(identical(.result,127)) {
  pass("Great job!")
} else {
  fail("Start over again!")
})
```

### 5.

<b>If you run</b>

```{r,eval = FALSE,echo = TRUE}
debug(ls)
```

<b>what happens when you next call the 'ls' function?</b>

```{r q5}
question('',
  answer("Execution of 'ls' will suspend at the beginning of the function and
         you will be in the browser.",correct = TRUE),
  answer("The 'ls' function will execute as usual."),
  answer("The 'ls' function will return an error."),
  answer("You will be prompted to specify at which line of the function you
         would like to suspend execution and enter the browser."))
```